# Week 11 Exercises

_McKinney 7.2, 11.1_

The general activity that we are doing in this week's exercise is to explore and try to understand a particular data set.  In this case, it is 

1. Read in the data file and filter down to only looking at MO hospitals
2. Aggregate by Hospital: sum the Denominator, use min Start Date, use max End Date  (watch for rows with no Denominator value!)
3. Calculate the average per day across that entire span
4. Histogram
5. Rank and find the hospital with the most

## STEP 1 - Read and Filter

<img src="images/step1.png" alt="Read and Filter Output" style="width: 500px; float: right; margin-left: 20px; border: 1px solid">

In the first step, read in the data file from the data directory `/data/complications_all.csv`.  It is a CSV file and Pandas should read it in just fine.  Explore the file so that you understand the columns and values.  At the end of this step, create a variable called `mo_hospitals` that contains a data frame from the `complications_all.csv` file, filtered down to only contain those hospitals from the state of Missouri (MO).

A screenshot is included for reference.



```python

```

<!-- grade_id: q1-solution, points: 0, grade: false, solution: true -->
```python
import pandas as pd

# This is just to show you the name to use for the variable you need to create for this step to pass.
x1 = pd.read_csv('complications_all.csv')

mo_hospitals = pd.DataFrame(x1[x1['State']=='MO'],columns=['Facility ID','Facility Name','Address','City','State','County Name','Zip Code','Phone Number','Measure ID',
                           'Meaure Name','Compared to National','Denominator','Score','Lower Estimate','Higher Estimate','Footnote','Start Date','End Date'])
print(mo_hospitals)

# Put your code below and make sure that you reassign `mo_hospitals` 
# to have the contents described in the instructions.


```

Output:
```
Facility ID                    Facility Name           Address  \
45534      260001            MERCY HOSPITAL JOPLIN     100 MERCY WAY   
45535      260001            MERCY HOSPITAL JOPLIN     100 MERCY WAY   
45536      260001            MERCY HOSPITAL JOPLIN     100 MERCY WAY   
45537      260001            MERCY HOSPITAL JOPLIN     100 MERCY WAY   
45538      260001            MERCY HOSPITAL JOPLIN     100 MERCY WAY   
...           ...                              ...               ...   
47662      263304  SHRINERS HOSPITALS FOR CHILDREN  4400 CLAYTON AVE   
47663      263304  SHRINERS HOSPITALS FOR CHILDREN  4400 CLAYTON AVE   
47664      263304  SHRINERS HOSPITALS FOR CHILDREN  4400 CLAYTON AVE   
47665      263304  SHRINERS HOSPITALS FOR CHILDREN  4400 CLAYTON AVE   
47666      263304  SHRINERS HOSPITALS FOR CHILDREN  4400 CLAYTON AVE   

              City State     County Name  Zip Code    Phone Number  \
45534       JOPLIN    MO          JASPER       NaN  (417) 781-2727   
4
```

```python
mo_hospitals.head()
```

Output:
```
Facility ID          Facility Name        Address    City State  \
45534      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45535      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45536      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45537      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   
45538      260001  MERCY HOSPITAL JOPLIN  100 MERCY WAY  JOPLIN    MO   

      County Name  Zip Code    Phone Number     Measure ID  Meaure Name  \
45534      JASPER       NaN  (417) 781-2727  COMP_HIP_KNEE          NaN   
45535      JASPER       NaN  (417) 781-2727    MORT_30_AMI          NaN   
45536      JASPER       NaN  (417) 781-2727   MORT_30_CABG          NaN   
45537      JASPER       NaN  (417) 781-2727   MORT_30_COPD          NaN   
45538      JASPER       NaN  (417) 781-2727     MORT_30_HF          NaN   

                      Compared to National Denominator Score Lower Estimate  \
45534  No Different Than the National
```

<!-- grade_id: q1-tests, points: 2, grade: true, solution: false -->
```python
# These assertions will help make sure that you're on the right track.
assert(mo_hospitals['State'].unique() == ['MO'])
assert(mo_hospitals.shape == (2133,18))
```

## STEP 2 - Transform and Aggregate

<img src="images/step2.png" alt="Transforma and Aggregate Output" style="width: 500px; float: right; margin-left: 20px; border: 1px solid">

In the next step, we need to aggregate the results by hospital.  There are some key fields that we want to summarize, though:
* We want to know the earliest date that each hospital was participating in any program
* We want to know the latest date that each hospital stopped participating in any program
* We want to know the total number of patients in the denominators of these programs

Some things to note:
* You will need to convert the `Start Date` and `End Date` to actual datetime fields
* You will need to clean up and convert the `Denominator` field to just be numeric - the rule that you should use it to simply remove any records where the `Denominator` is `'Not Available'`

The final result of this step should be a new data frame called `mo_summary` that contains one row for each hospital and contains the min start date, max end date, and total denominator.  Use the names `start_date`, `end_date`, and `number` for those columns in `mo_summary`.

A screenshot is included for reference.

<!-- grade_id: q2-solution, points: 0, grade: false, solution: true -->
```python
# This is just to show you the name to use for the variable you need to create for this step to pass.
import pandas as pd
mo_summary = pd.DataFrame()

mo_hospitals['Start Date'] = pd.to_datetime(mo_hospitals['Start Date'])
mo_hospitals['End Date'] = pd.to_datetime(mo_hospitals['End Date'])
mo_hospitals = mo_hospitals[mo_hospitals['Denominator']!='Not Available']
mo_hospitals['Facility Name'].unique()

#mo_hospitals.groupby('Facility Name')['Denominator'].unique()
#ertz = pd.DataFrame(mo_hospitals.groupby('Facility Name')['Denominator'].unique())
#Could not get pandas to sum each list as their own element in number
#number=ertz.sum()
#ertz, that is above, is the best display of the denominator values I could get, but I could not get them to sum the list for each unique facility name.

number=mo_hospitals.groupby('Facility Name')['Denominator'].sum().unique()
#could not get number to acutally sum the values, only to smash them into one large number
number = pd.Series(number)
number = number.rename("Averages")
#Also could not get number to be under the facility names columns



end_date = mo_hospitals.groupby('Facility Name')['End Date'].max()
start_date = mo_hospitals.groupby('Facility Name')['Start Date'].min()

mo_summary = mo_summary.append(start_date)
mo_summary = mo_summary.append(end_date)
mo_summary = mo_summary.append(number)
#Could not get it to work -> mo_summary.set_index(mo_hospitals['Facility Name'].unique())
mo_summary

# Put your code below and make sure that you reassign `mo_summary` 
# to have the contents described in the instructions.


```

Output:
```
BARNES JEWISH HOSPITAL BARNES-JEWISH ST PETERS HOSPITAL  \
Start Date             2015-04-01                       2015-04-01   
End Date               2018-06-30                       2018-06-30   
Averages                      NaT                              NaT   

           BARNES-JEWISH WEST COUNTY HOSPITAL BATES COUNTY MEMORIAL HOSPITAL  \
Start Date                         2015-04-01                     2015-07-01   
End Date                           2018-06-30                     2018-06-30   
Averages                                  NaT                            NaT   

           BELTON REGIONAL MEDICAL CENTER BOONE HOSPITAL CENTER  \
Start Date                     2015-04-01            2015-04-01   
End Date                       2018-06-30            2018-06-30   
Averages                              NaT                   NaT   

           BOTHWELL REGIONAL HEALTH CENTER CAMERON REGIONAL MEDICAL CENTER  \
Start Date                      2015-04-01                    
```

<!-- grade_id: q2-test, points: 5, grade: true, solution: false -->
```python
assert(mo_summary['number'].sum() == 1766908)
assert(mo_summary['start_date'].min() == pd.Timestamp(2015,4,1))
assert(mo_summary['end_date'].max() == pd.Timestamp(2018,6,30))
assert(mo_summary.shape == (108,3))
assert(mo_summary.loc['BARNES JEWISH HOSPITAL'].number == 131313)
assert(mo_summary.loc['BOONE HOSPITAL CENTER'].number == 63099)
```

## STEP 3: Average Per Day

<img src="images/step3.png" alt="Average per Day" style="width: 500px; margin-left: 20px; float: right; border: 1px solid">


The next step, now that we have a start date, end date, and total patients for each day is to calculate how many patients on average per day this represents.  You will need to calculate the number of days between the start date and end date, and then the average as total patients divided by number of days.

Your final output should still be the variable `mo_summary`, and it will need to have two additional columns: `days` and `per_day`.

A screenshot is included for reference.


<!-- grade_id: q3-solution, points: 0, grade: false, solution: true -->
```python
# Put your code below and make sure that you add new columns to `mo_summary` 
# to have the contents described in the instructions.
days = end_date-start_date
#Because number would not go under the facility names as columns, per_day will not work
per_day = number/days
mo_summary.append(days,ignore_index=True)
mo_summary.append(per_day,ignore_index=True)

```

Output:
```
/opt/tljh/user/lib/python3.6/site-packages/pandas/core/arrays/timedeltas.py in __rtruediv__(self, other)
    584             #  inference on the result.  It does not raise, a numeric array
    585             #  is returned.  GH#23829
--> 586             result = [other[n] / self[n] for n in range(len(self))]
    587             return np.array(result)
    588 

/opt/tljh/user/lib/python3.6/site-packages/pandas/core/arrays/timedeltas.py in <listcomp>(.0)
    584             #  inference on the result.  It does not raise, a numeric array
    585             #  is returned.  GH#23829
--> 586             result = [other[n] / self[n] for n in range(len(self))]
    587             return np.array(result)
    588 

TypeError: unsupported operand type(s) for /: 'str' and 'NaTType'
```

```python

```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert(mo_summary['days'].sum() == 123718)
assert(mo_summary['per_day'].mean() == 13.841279932122697)
assert(mo_summary['per_day'].min() == 0.023744292237442923)
assert(mo_summary['days'].min() == 729)
assert(mo_summary['days'].max() == 1186)
```

## STEP 4: Histogram

<img src="images/step4.png" alt="Histogram" style="width: 300px; float: right; margin-left: 20px; border: 1px solid">

For this step, I want you to plot a histogram to give yourself a better understanding of this new metric we've calculated: how many patients per day have a potential complication.  Looking at the shape and distribution of the data might give us some indication of the hospital volume and quality.  (Don't take this too seriously, though.  We're cutting some important corners.)

A screenshot is included for reference.

<!-- grade_id: q4-solution, points: 3, grade: true, solution: true -->
```python
%matplotlib inline
# Use the built-in Pandas histogram plotting capability to plot a histogram of the `per_day` values.
import pandas as pd
mo_summary['avg'].hist(bins=5)


```

Output:
```
pandas/_libs/hashtable_class_helper.pxi in pandas._libs.hashtable.PyObjectHashTable.get_item()

pandas/_libs/hashtable_class_helper.pxi in pandas._libs.hashtable.PyObjectHashTable.get_item()

KeyError: 'avg'
```

## STEP 5: Rank Order

Now that we see what the distribution looks like, you can tell that most of the `per_day` values are below 20.  There are, however, some outliers up around 100!  Who are those?  Let's rank the data set by the `per_day` value in descending order, examine them, and produce a list of the "top 3" based on this criteria.  That is, who are the top three hospitals based on having the highest `per_day` values.

At the end of this step, put those hospital names in a list called `top_hospitals` and that will be used for testing.  They should be in the order they appear `[#1, #2, #3]`.  That is, if the top hospitals were Mercy (99.3), BJC (97.2), and MoBap (90.1), then you would have `top_hospitals = ['Mercy','BJC','MoBap']`

<!-- grade_id: q5-solution, points: 0, grade: false, solution: true -->
```python
# Put your code below and make sure that your final result ends up in the variable `top_hospitals`
# as described in the instructions above.
import pandas as pd
top_hospitals = []
top_hospitals = mo_summary[3].all().rank()



```

Output:
```
AttributeError                            Traceback (most recent call last)
<ipython-input-12-04176ffaa794> in <module>
      3 import pandas as pd
      4 top_hospitals = []
----> 5 top_hospitals = mo_summary[3].all().rank()
      6 
      7 

AttributeError: 'numpy.ndarray' object has no attribute 'rank'
```

<!-- grade_id: q5-tests, points: 5, grade: true, solution: false -->
```python
assert(type(top_hospitals) == list)
### BEGIN HIDDEN TESTS
assert(top_hospitals == ['BARNES JEWISH HOSPITAL', 'MERCY HOSPITAL SPRINGFIELD', 'COX MEDICAL CENTERS'])
### END HIDDEN TESTS
```

```python
#this is all i could get done
```